<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 0  -->
  <question type="category">
    <category>
      <text>$course$/top/2020年度夏学期・月1月2・プログラミング演習Ⅰ（B）（亀井　靖高） のデフォルト/演習課題2の小テスト</text>
    </category>
    <info format="html">
      <text><![CDATA[<p dir="ltr" style="text-align: left;">演習課題2に対する小テストの集合です<br></p>]]></text>
    </info>
    <idnumber></idnumber>
  </question>

<!-- question: 9635  -->
  <question type="coderunner">
    <name>
      <text>04-Loop.pdfの課題10</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p dir="ltr" style="text-align: left;"><div><div>半角文字からなる文字列を読み込んで，逆転させた文字列をつないで回文を作り，それを表示するプログラムを作ってください．</div></div><br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>c_program</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[#include <stdio.h>
#include <string.h>
int main(){
    char x[100], y[100];
    int i;
    scanf("%s", x);
    for(i=0; i<strlen(x); i++){
       y[i]=x[strlen(x)-i-1]; 
    }
    y[strlen(x)]='\0';
    strcat(x,y);
    printf("%s",x);
}]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>04-10-test1</text>
      </testcode>
      <stdin>
                <text>aiueo</text>
      </stdin>
      <expected>
                <text>aiueooeuia
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>04-10-test2</text>
      </testcode>
      <stdin>
                <text>abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVW</text>
      </stdin>
      <expected>
                <text>abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWWVUTSRQPONMLKJIHGFEDCBAzyxwvutsrqponmlkjihgfedcba
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>04-10-test3</text>
      </testcode>
      <stdin>
                <text>abcdefghijklmnopqrstuvwxyz</text>
      </stdin>
      <expected>
                <text>abcdefghijklmnopqrstuvwxyzzyxwvutsrqponmlkjihgfedcba</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>04-10-test4</text>
      </testcode>
      <stdin>
                <text>aaaa</text>
      </stdin>
      <expected>
                <text>aaaaaaaa</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>04-10-test5</text>
      </testcode>
      <stdin>
                <text>
</text>
      </stdin>
      <expected>
                <text>
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 9633  -->
  <question type="coderunner">
    <name>
      <text>04-Loop.pdfの課題8.1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p dir="ltr" style="text-align: left;">04-Loop.pdfのベクトルの内積を計算するプログラムを書き換えて，キーボードからの入力を読み込むようにせよ．<br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>c_program</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[#include <stdio.h>
#include <math.h>

int main()
{
    int i;
    double ip = 0.0;
    double u[3];
    double v[3];
    
    
    for(i=0; i<3; i++){
        scanf("%lf", &u[i]);
    }
    
        
    for(i=0; i<3; i++){
        scanf("%lf", &v[i]);
    }
    
    for (i=0; i<3; i++){
        ip = ip + u[i]*v[i];
    }
    printf("%f\n", ip);
    
    return 0;
}]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>04-8-1-test1</text>
      </testcode>
      <stdin>
                <text>1
2
3
4
5
6</text>
      </stdin>
      <expected>
                <text>32.000000
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>04-8-1-test2</text>
      </testcode>
      <stdin>
                <text>3
-6
-4
6
9
-9</text>
      </stdin>
      <expected>
                <text>0.000000
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>04-8-1-test3</text>
      </testcode>
      <stdin>
                <text>1.2
-2.7
-4.5
-2.0
-1.6
-2.3</text>
      </stdin>
      <expected>
                <text>12.270000
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>04-8-1-test4</text>
      </testcode>
      <stdin>
                <text>0
0
0
0
0
0</text>
      </stdin>
      <expected>
                <text>0.000000
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>04-8-1-test5</text>
      </testcode>
      <stdin>
                <text>135
278
984
-234
-126
-332</text>
      </stdin>
      <expected>
                <text>-393306.000000
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 9634  -->
  <question type="coderunner">
    <name>
      <text>04-Loop.pdfの課題8.2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p dir="ltr" style="text-align: left;"><div><div>3×3の行列の積を計算するプログラムを作成して下さい．</div></div><br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>c_program</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[#include <stdio.h>
#include <math.h>

int main()
{
    int i, j, k;
    float answer;
    float a[3][3], b[3][3], c[3][3];
    
    for(i=0; i<3; i++){
        for(j=0; j<3; j++){
                scanf("%f", &a[i][j]);
        }
    }
    
    for(i=0; i<3; ++i){
        for(j=0; j<3; ++j){
                scanf("%f", &b[i][j]);
        }
    }
    
    //行列の積
    for(i=0; i<3; ++i){
        for(j=0; j<3; ++j){
            c[i][j] = a[i][j]*b[i][j];
            
            for(i=0; i<3; ++i){
                for(j=0; j<3; ++j){
                    answer = 0;
                    for (k=0; k<3; ++k){
                        answer = answer + a[i][k]*b[k][j];
                        c[i][j] = answer;
                    }
                }
            }
        }
    }
    
    for(i=0; i<3; ++i){
        for(j=0; j<3; ++j){
            printf("%f\n",c[i][j]);
        }
    }
    
    return 0;
}]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>04-8-2-test1</text>
      </testcode>
      <stdin>
                <text>1
2
3
2
3
2
3
2
1
3
2
1
1
3
2
2
1
3</text>
      </stdin>
      <expected>
                <text>11.000000
11.000000
14.000000
13.000000
15.000000
14.000000
13.000000
13.000000
10.000000
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>04-8-2-test2</text>
      </testcode>
      <stdin>
                <text>1
0
0
0
1
0
0
0
1
1
0
0
0
1
0
0
0
1</text>
      </stdin>
      <expected>
                <text>1.000000
0.000000
0.000000
0.000000
1.000000
0.000000
0.000000
0.000000
1.000000
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>04-8-2-test3</text>
      </testcode>
      <stdin>
                <text>-1.5
2.7
3.5
1.2
1.0
-4.5
0
0
0
-3.4
2.3
5.4
9.4
0
-2.1
-2.5
-3.2
1.3</text>
      </stdin>
      <expected>
                <text>21.730000
-14.650000
-9.220001
16.570000
17.160000
-1.469999
0.000000
0.000000
0.000000
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>04-8-2-test4</text>
      </testcode>
      <stdin>
                <text>24
46
-32
-11
-34
-12
64
73
-49
-23
-45
-26
14
87
91
-59
-32
-18</text>
      </stdin>
      <expected>
                <text>1980.000000
3946.000000
4138.000000
485.000000
-2079.000000
-2592.000000
2441.000000
5039.000000
5861.000000
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>04-8-2-test5</text>
      </testcode>
      <stdin>
                <text>1
0
-1
0
-1
1
-1
1
0
1
0
-1
0
-1
1
-1
1
0</text>
      </stdin>
      <expected>
                <text>2.000000
-1.000000
-1.000000
-1.000000
2.000000
-1.000000
-1.000000
-1.000000
2.000000
</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

</quiz>